# Assignment 4
In this assignment you will be using the dataset released by The Department of Transportation. This dataset lists flights that occurred in 2015, along with other information such as delays, flight time etc.

In this assignment, you will be showing good practices to manipulate data using Python's most popular libraries to accomplish the following:

- cleaning data with pandas
- make specific changes with numpy
- handling date-related values with datetime

Note: please consider the flights departing from BOS, JFK, SFO and LAX.

Each question is equally weighted for the total grade.

```python
import os
import pandas as pd
import pandas.api.types as ptypes
import numpy as np
import datetime as dt
from mads.lib.path import assets

# Retrieve data
file_flights = assets.find("flights.csv")
file_airports = assets.find("airports.csv")
file_airlines = assets.find("airlines.csv")
flights_df_raw = pd.read_csv(file_flights, low_memory=False)
airports_df = pd.read_csv(file_airports)
airlines_df = pd.read_csv(file_airlines)
```

## Question 1: Data Preprocessing
For this question, perform the following:

- remove rows with missing values
- keep flights departing from airports (ORIGIN_AIRPORT) that we want to look at (BOS, JFK, SFO and LAX)
- filter out the flights that have more than 1 day delay (DEPARTURE_DELAY)
- convert FLIGHT_NUMBER column type to string
- SCHEDULED_DEPARTURE is coded as a float where the first two digits indicate the hour and the last two indicate the minutes. Convert this column to datetime format by combining existing columns DAY, MONTH, YEAR and SCHEDULED_DEPARTURE
- add IS_DELAYED column by considering any flight with at least 15 minutes delay (DEPARTURE_DELAY) are delayed, and any other flight is not delayed
- remove YEAR, MONTH, DAY columns

```python
flights_df_raw[flights_df_raw["DEPARTURE_DELAY"] > 1440].head()
```

Output:
```
YEAR  MONTH  DAY ORIGIN_AIRPORT DESTINATION_AIRPORT AIRLINE  \
145439   2015      1   18            LAS                 LAX      AA   
1225250  2015      5   21            LAX                 STL      AA   
1518916  2015      6   22            SFO                 DFW      AA   
2567756  2015     11   15            MCO                 JFK      AA   
2667908  2015     11   27            DTW                 ORD      AA   

         FLIGHT_NUMBER  SCHEDULED_DEPARTURE  DEPARTURE_DELAY  
145439             224                 1130           1604.0  
1225250           1258                  610           1492.0  
1518916           1454                  515           1496.0  
2567756            290                 1215           1536.0  
2667908           2559                 1027           1631.0
```

<!-- grade_id: cell-85d6b39dfr7aa02f, points: 0, grade: false, solution: true -->
```python
def data_preprocess(flights_df):
    # YOUR CODE HERE
    df = flights_df.dropna()
    df = df[df['ORIGIN_AIRPORT'].isin(['BOS', 'JFK', 'SFO', 'LAX'])]
    df = df[df['DEPARTURE_DELAY'] < 1440]
    df['FLIGHT_NUMBER'] = df['FLIGHT_NUMBER'].astype(str)
    #scheduled_departure
    df['SCHEDULED_DEPARTURE'] = pd.to_datetime(
    df[['YEAR', 'MONTH', 'DAY']].astype(str).agg('-'.join, axis=1) + ' ' +
    df['SCHEDULED_DEPARTURE'].astype(str).str.zfill(4),
    format='%Y-%m-%d %H%M')
    #
    df['IS_DELAYED'] = df['DEPARTURE_DELAY'].apply(lambda x: 1 if x >= 15 else 0)
    df = df.drop(columns=['YEAR', 'MONTH', 'DAY'])
    return df
```

```python
#flights_df = data_preprocess(flights_df_raw.copy())
#flights_df.head()
```

<!-- grade_id: cell-ece042ccds21104a, points: 1, grade: true, solution: false -->
```python
flights_df = data_preprocess(flights_df_raw.copy())
assert len(flights_df) == 535744, "Q1: There should be 535744 observations in the flights dataframe"

```

```python
flights_df['FLIGHT_N
```

## Question 2
Merge flights_df dataframe with airports_df dataframe and return the number of departing flights (*NUM_FLIGHTS*) per airport (*IATA_CODE*) across the year.

```python
#flights_df.head()
```

```python
#airports_df.head()
```

```python
#flights_df['ORIGIN_AIRPORT'].unique() 
```

<!-- grade_id: cell-bce4d6f2ecddh297, points: 0, grade: false, solution: true -->
```python
def flights_per_airport(flights_df, airports_df):
    flights_df = data_preprocess(flights_df.copy())
    merged_df = pd.merge(flights_df, airports_df, how='left', left_on='ORIGIN_AIRPORT', right_on='IATA_CODE')
    #print(merged_df.shape)
    #print(merged_df['ORIGIN_AIRPORT'].unique()) 
    df = merged_df.groupby('IATA_CODE').size().reset_index(name='NUM_FLIGHTS').set_index('IATA_CODE')

    #print(df.shape)
    return df
```

```python
#num_flights_df=flights_per_airport(flights_df_raw.copy(), airports_df.copy())
```

```python
#num_flights_df.head()
```

<!-- grade_id: cell-780b5a4da845dbce, points: 1, grade: true, solution: false -->
```python
num_flights_df=flights_per_airport(flights_df_raw.copy(), airports_df.copy())

assert num_flights_df.shape==(4,1), "Shape of DataFrame should be (4,1)"
assert num_flights_df.columns[0]=='NUM_FLIGHTS', "DataFrame should have a column which is called NUM_FLIGHTS"
assert num_flights_df.loc["BOS", "NUM_FLIGHTS"] == 105276, "The NUM_FLIGHTS for BOS is wrong"

```

## Question 3
For this question, find the top three airline names which have high number of flights and the least percentage of delay compared to other airlines. The result should be a dataframe which has three columns *AIRLINE_NAME*, *NUM_FLIGHTS* and *PERC_DELAY*.

Hint:
- percentage of delay for each airline is obtained using groupby and apply methods
- merge flights_df with airlines_df to get the names of top three airlines

```python
#num_flights_df.head()
```

```python
#flights_df_raw.head()
```

```python
#airlines_df.head()
```

```python
#airlines_df_copy = airlines_df.rename(columns={'AIRLINE': 'AIRLINE_NAME'})
#airlines_df_copy.head()
```

```python
#merged_df = pd.merge(data_preprocess(flights_df_raw.copy()), airlines_df_copy.copy(), how='left', left_on='AIRLINE' , right_on='IATA_CODE')
#df = merged_df.groupby('AIRLINE').size().reset_index(name='NUM_FLIGHTS').set_index('AIRLINE')

#df = (merged_df.groupby('AIRLINE')
#      .agg(NUM_FLIGHTS=('AIRLINE', 'count')
#           ,PERC_DELAY=('IS_DELAYED', 'mean')
#           ,AIRLINE_NAME=('AIRLINE_NAME', 'first'))
#      .rename_axis('AIRLINE'))
#
#df = df.sort_values(by=['NUM_FLIGHTS', 'PERC_DELAY', 'AIRLINE_NAME'], ascending=[False, False, True])
#
#df = df.head(3)
```

<!-- grade_id: cell-f94ca014da9kaf38, points: 0, grade: false, solution: true -->
```python
def top_three_airlines(flights_df, airlines_df):
    # YOUR CODE HERE
    airlines_df_copy = airlines_df.rename(columns={'AIRLINE': 'AIRLINE_NAME'})

    merged_df = pd.merge(data_preprocess(flights_df), airlines_df_copy, how='left', left_on='AIRLINE' , right_on='IATA_CODE')

    df = (merged_df.groupby('AIRLINE')
          .agg(NUM_FLIGHTS=('AIRLINE', 'count')
               ,PERC_DELAY=('IS_DELAYED', 'mean')
               ,AIRLINE_NAME=('AIRLINE_NAME', 'first'))
          .reset_index()
          .drop(columns='AIRLINE'))

    df = df.sort_values(by=['NUM_FLIGHTS', 'PERC_DELAY', 'AIRLINE_NAME'], ascending=[False, False, True])

    df = df.head(3).reset_index(drop=True)
    return df
```

```python
#top_three_airlines_df = top_three_airlines(flights_df_raw.copy(), airlines_df.copy())
#top_three_airlines_df.head()
#print(top_three_airlines_df.loc[0, 'AIRLINE_NAME'])
```

<!-- grade_id: cell-3c64ee33f975w296, points: 1, grade: true, solution: false -->
```python
top_three_airlines_df = top_three_airlines(flights_df_raw.copy(), airlines_df.copy())

assert sorted(list(top_three_airlines_df.columns)) == sorted(['NUM_FLIGHTS', 'PERC_DELAY', 'AIRLINE_NAME']), "Dataframe doesn't have required columns"
assert top_three_airlines_df.loc[0, 'AIRLINE_NAME'] == 'United Air Lines Inc.', "Top airline name doesn't match"

```

## Question 4
For this question, obtain the monthly percentage of delays for each *ORIGIN_AIRPORT*.

Example Result:

         MONTH     BOS     JFK     LAX     SFO
    0   January  0.1902  0.2257  0.1738  0.xxxx
    1  February  0.3248  0.xxxx  0.xxxx  0.xxxx
    2     March  0.1984  0.xxxx  0.xxxx  0.xxxx
    3     April  0.xxxx  0.xxxx  0.xxxx  0.xxxx

```python
#df = data_preprocess(flights_df_raw.copy())
#df.head()
```

```python
#flights_df['MONTH'].unique()
```

```python
#df.shape
```

<!-- grade_id: cell-9f15bd32a4tea83d, points: 0, grade: false, solution: true -->
```python
def monthly_airport_delays(flights_df):
    # YOUR CODE HERE
    flights_df = data_preprocess(flights_df)
    flights_df['MONTH'] = flights_df['SCHEDULED_DEPARTURE'].dt.strftime('%B')

    month_order = ['January','February','March','April','May','June',
               'July','August','September','October','November','December']
    flights_df['MONTH'] = pd.Categorical(flights_df['MONTH'], categories=month_order, ordered=True)

    df = flights_df.pivot_table(index='MONTH', columns='ORIGIN_AIRPORT', values='IS_DELAYED', aggfunc='mean')
    #print(df)
    df.columns.name = None
    df = df.reset_index()
    df['MONTH'] = df['MONTH'].astype(str)
    df = df.round(4)
    return df
```

<!-- grade_id: cell-8c8d2b95b2b8c3eq, points: 1, grade: true, solution: false -->
```python
monthly_airport_delays_df = monthly_airport_delays(flights_df_raw.copy())

```

```python
#monthly_airport_delays_df
```