# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [NISPUF17.csv](NISPUF17.csv). A data users guide for this, which you'll need to understand and map the variables in the data to the questions being asked, is available at [NIS-PUF17-DUG.pdf](NIS-PUF17-DUG.pdf) and [NIS-PUF17-CODEBOOK.pdf](NIS-PUF17-CODEBOOK.pdf). 

**Note: For all assignments please write all of your code within the function we define in order to ensure it is run by the autograder correctly**

## Question 1: proportion_of_education()
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
import pandas as pd
```

<!-- grade_id: cell-eea16d020eb52ae6, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    # column EDUC1:
    # less than high school = 1
    # high school = 2
    # more than high school but not college = 3
    # college = 4
    
    file = open("NISPUF17.csv", "r")
    df = pd.read_csv(file)
    #df = pd.read_csv("NISPUF17.csv")
    
    total_children = df["CHILDNM"].count()
    less_than_high_school = (df[df["EDUC1"] == 1]["CHILDNM"].count() / total_children)
    high_school = (df[df["EDUC1"] == 2]["CHILDNM"].count() / total_children)
    more_than_high_school_but_not_college = (df[df["EDUC1"] == 3]["CHILDNM"].count() / total_children)
    college = (df[df["EDUC1"] == 4]["CHILDNM"].count() / total_children)
    Total_percentage = less_than_high_school + high_school + more_than_high_school_but_not_college + college
    #print(f"""
    #Total children: {total_children}
    #Less than high school: {less_than_high_school}
    #High school: {high_school}
    #More than high school but not college: {more_than_high_school_but_not_college}
    #College: {college}
    #Total percentage = {Total_percentage}
    #""")
    d = {"less than high school": less_than_high_school,
    "high school": high_school,
    "more than high school but not college": more_than_high_school_but_not_college,
    "college": college}   
    return d
    
    #raise NotImplementedError()
```

```python
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51d, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2: chickenpox_by_sex()
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.00779`.

<!-- grade_id: cell-a0a9e6fe67698002, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    # SEX - 1 = MALE | 2 = FEMALE
    #P_NUMVRC - total number of varicella doses
    #HAD_CPOX - 1 = YES | 2 = NO

    file = open("NISPUF17.csv", "r")
    df = pd.read_csv(file)
    #print(df[["SEX", "P_NUMVRC", "HAD_CPOX"]].head())
    
    #total_children = df["CHILDNM"].count()

    #Calculate male
    total_children_male_contracted = ((df["SEX"] == 1) & (df["HAD_CPOX"] == 1)).sum()
    total_children_male_vax_contract = ((df["SEX"] == 1) & (df["HAD_CPOX"] == 1) & (df["P_NUMVRC"] >= 1)).sum()
    total_children_male_vax_non_contract = ((df["SEX"] == 1) & (df["HAD_CPOX"] == 2) & (df["P_NUMVRC"] >= 1)).sum()

    ratio_male = total_children_male_vax_contract / total_children_male_vax_non_contract

    #Calculate female
    total_children_female_contracted = ((df["SEX"] == 2) & (df["HAD_CPOX"] == 1)).sum()
    total_children_female_vax_contract = ((df["SEX"] == 2) & (df["HAD_CPOX"] == 1) & (df["P_NUMVRC"] >= 1)).sum()
    total_children_female_vax_non_contract = ((df["SEX"] == 2) & (df["HAD_CPOX"] == 2) & (df["P_NUMVRC"] >= 1)).sum()

    ratio_female = total_children_female_vax_contract / total_children_female_vax_non_contract
    

    return {
        "male": ratio_male
        ,"female": ratio_female
    }
```

```python
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c864, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 3: corr_chickenpox()
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no's) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses. 

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79b, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #               "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    file = open("NISPUF17.csv", "r")
    df = pd.read_csv(file)
    df_corr = df[["HAD_CPOX", "P_NUMVRC"]].dropna()
    df_corr = df_corr[df_corr["HAD_CPOX"] != 77]
    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    corr, pval=stats.pearsonr(df_corr["HAD_CPOX"], df_corr["P_NUMVRC"])

    print(df_corr.groupby("P_NUMVRC")["HAD_CPOX"].mean())
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
#MAx - alternate solution - filter out children who contracted before they could get vaccinated (0-12 month contracted)
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    file = open("NISPUF17.csv", "r")
    df = pd.read_csv(file)
    df_corr = df[["HAD_CPOX", "P_NUMVRC", "AGECPOXR"]].dropna(subset=["HAD_CPOX", "P_NUMVRC"])
    df_corr = df_corr[df_corr["HAD_CPOX"] != 77]
    df_corr = df_corr[df_corr["AGECPOXR"].isin([3, 4]) | df_corr["AGECPOXR"].isna()]
    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df_corr["HAD_CPOX"], df_corr["P_NUMVRC"])

    print(df_corr.groupby("P_NUMVRC")["HAD_CPOX"].mean())
    
    # just return the correlation
    return corr


```

```python
corr_chickenpox()
```

Output:
```
P_NUMVRC
0.0    1.961049
1.0    1.991483
2.0    1.981221
3.0    2.000000
Name: HAD_CPOX, dtype: float64
0.0704487346014799
```

```python

counts_dict = {
    "Zero dose": (df["P_NUMVRC"] == 0).sum()
    ,"One dose": (df["P_NUMVRC"] == 1).sum()
    ,"Two dose": (df["P_NUMVRC"] == 2).sum()
    ,"Three dose": (df["P_NUMVRC"] == 3).sum()
}

counts_dict
```

Output:
```
{'Zero dose': 1338, 'One dose': 13781, 'Two dose': 213, 'Three dose': 1}
```

```python
#Trying other stuff to explain dose 2 anomaly

import scipy.stats as stats
import numpy as np
import pandas as pd

file = open("NISPUF17.csv", "r")
df = pd.read_csv(file)

def age_to_bucket(age):
    if age <= 6: return 1
    elif age <= 12: return 2
    elif age <= 18: return 3
    else: return 4

df["VRC1_AGE_BUCKET"] = df["VRC1_AGE"].apply(age_to_bucket)
df["VRC2_AGE_BUCKET"] = df["VRC2_AGE"].apply(age_to_bucket)

#df = df[["AGECPOXR", "VRC1_AGE_BUCKET", "VRC2_AGE_BUCKET", "P_NUMVRC"]].dropna()

#childs_who_had_second_dose = len(df)
#childs_contracted_after_second_dose = len(df[df["AGECPOXR"] > df["VRC2_AGE_BUCKET"]])
#childs_contracted_before_first_dose = len(df[df["AGECPOXR"] < df["VRC1_AGE_BUCKET"]])
childs_who_had_1_doses = len(df[df["P_NUMVRC"] == 1])
childs_who_had_1_doses_and_contracted = len(df[(df["P_NUMVRC"] == 1) & (df["HAD_CPOX"] == 1)])
childs_who_had_1_doses_and_not_contracted = len(df[(df["P_NUMVRC"] == 1) & (df["HAD_CPOX"] == 2)])
print(childs_who_had_1_doses, childs_who_had_1_doses_and_contracted, childs_who_had_1_doses_and_not_contracted)
print(childs_who_had_1_doses_and_not_contracted / childs_who_had_1_doses)


childs_who_had_2_doses = len(df[df["P_NUMVRC"] == 2])
childs_who_had_2_doses_and_contracted = len(df[(df["P_NUMVRC"] == 2) & (df["HAD_CPOX"] == 1)])
childs_who_had_2_doses_and_not_contracted = len(df[(df["P_NUMVRC"] == 2) & (df["HAD_CPOX"] == 2)])
print(childs_who_had_2_doses, childs_who_had_2_doses_and_contracted, childs_who_had_2_doses_and_not_contracted)
print(childs_who_had_2_doses_and_not_contracted / childs_who_had_2_doses)

df[(df["P_NUMVRC"] == 2) & (df["HAD_CPOX"] == 1)][["AGECPOXR", "VRC1_AGE_BUCKET", "VRC2_AGE_BUCKET"]]
#childs_contracted_before_first_dose_and_got_2_doses = len(df[(df["AGECPOXR"] < df["VRC1_AGE_BUCKET"]) & (df["P_NUMVRC"] == 2)])

#r_contracted_and_got_second_dose = childs_contracted_before_first_dose_and_got_2_doses / childs_contracted_before_first_dose



#df_corr = df_corr[["HAD_CPOX", "P_NUMVRC"]].dropna()
#df_corr = df_corr[df_corr["HAD_CPOX"] != 77]
#df_corr["HAD_CPOX"] = df_corr["HAD_CPOX"].replace({1: 1, 2: 0})

#print(df_corr.groupby("P_NUMVRC")["HAD_CPOX"].mean())
```

Output:
```
13781 117 13620
0.9883172483854582
213 4 209
0.9812206572769953
       AGECPOXR  VRC1_AGE_BUCKET  VRC2_AGE_BUCKET
5613        3.0                4                4
6110        4.0                2                3
9092        3.0                3                4
14156       4.0                3                4
```

<!-- grade_id: cell-73408733543a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
P_NUMVRC
0.0    1.961049
1.0    1.991483
2.0    1.981221
3.0    2.000000
Name: HAD_CPOX, dtype: float64
```

Max - Yes as shown above there are fields you can use to evaluate when children are contracting vs when they were dosed. The dose two anomaly must just be down to low sample size (only 200 ish who got 2 doses so even 4 contracted results in higher percentage than dose one contracted)

```python
#more checking - likelyhood to contract cpox if vaxed vs not vaxed then ratio for how many times more likely to contract if not vaxed.
file = open("NISPUF17.csv", "r")
df = pd.read_csv(file)

df = df[["HAD_CPOX", "P_NUMVRC"]].dropna()

total_vax = (df["P_NUMVRC"] >= 1).sum()
total_non_vax = (df["P_NUMVRC"] == 0).sum()

vax_contracted = ((df["HAD_CPOX"] == 1) & (df["P_NUMVRC"] >= 1)).sum()
non_vax_contracted = ((df["HAD_CPOX"] == 1) & (df["P_NUMVRC"] == 0)).sum()

perc_vax_contracted = (vax_contracted / total_vax) * 100
perc_non_vax_contracted = (non_vax_contracted / total_non_vax) * 100

likely = perc_non_vax_contracted / perc_vax_contracted

print(perc_vax_contracted, perc_non_vax_contracted, likely)


```

Output:
```
0.8645944980350125 3.8863976083707024 4.495052440425453
```

```python

```